# Day 13 notebook

The objectives of this notebook are to practice

* representing trees as strings
* visualizing trees
* enumerating trees
* examining trees

## The `toytree` module

In this activity we will use the `toytree` (third-party) module to work with phylogenetic trees.  In particular, we will use the newick-format parsing and visualization functionality of this module.  I encourage you to look through the [`toytree` documentation](https://toytree.readthedocs.io/en/latest/4-tutorial.html).

```python
import toytree # for working with tree data structures
```

## The newick format for trees
The classic string representation for trees is the [Newick format](https://en.wikipedia.org/wiki/Newick_format).  In addition to the Wikipedia article, you can also read the [description by Joe Felsenstein](https://phylipweb.github.io/phylip/newicktree.html), one of the leaders in the field of phylogenetics (and also a UW-Madison alumnus!).

For example, the Newick formatted string `"((A, B),((C, D),E));"` corresponds to the tree below:

```python
example_newick = "((A, B), ((C, D), E));"
example_tree = toytree.tree(example_newick)
example_tree.draw(use_edge_lengths=False); # trailing semicolon to suppress printing of return value
```

You can specify branch lengths in a Newick string using a colon (`:`) followed by a number after the child node of the branch.  For example, the Newick formatted string `"((A:3, B:3):1,((C:1,D:1):2,E:3):1);"` corresponds to the tree below:

```python
example2_newick = "((A:3, B:3):1, ((C:1, D:1):2, E:3):1);"
example2_tree = toytree.tree(example2_newick)
example2_tree.draw(use_edge_lengths=True, scalebar=True);
```

Here is an alternative tree visualization of the same tree.  Note that the lengths of the vertical lines in the previous visualization are not meaningful.

```python
example2_tree.draw(tree_style="c", tip_labels=True);
```

The Newick format also allows for specifying the names and support values of internal nodes, but we will not generally use those features, so do not worry about that aspect of the format for now.

## PROBLEM 1: Newick string for a tree without branch lengths (1 POINT)

Give a Newick string (without branch lengths) that respresents the tree shown below.  Assign that string to the variable `problem1_newick` in the cell below.
![problem1_tree](problem1_tree.png)

```python
###
problem1_newick ="((E, F),(D,(C,(A,B))));"

#"((((B,A),C),D),(F,E))"
###

```

<!-- grade_id: problem1_newick, points: 1, grade: true, solution: false -->
```python
# tests for problem1_newick
problem1_tree = toytree.tree(problem1_newick) # this will raise an exception if the string is not valid
###
### AUTOGRADER TEST - DO NOT REMOVE
###
print("SUCCESS: problem1_newick passed all visible tests")
```

Output:
```
SUCCESS: problem1_newick passed all visible tests
```

## PROBLEM 2: Newick string for a tree with branch lengths (1 POINT)

Give a Newick string (with branch lengths) that respresents the tree shown below.  Assign that string to the variable `problem2_newick` in the cell below.
![problem2_tree](problem2_tree.png)

```python
###

problem2_newick = "((A:3, (B:2, C:2)):1, ((D:1, E:1), F:2):2);"
#"((A:3, B:2):1, ((C:1, D:1):2, E:3):1);"
###
example2_tree = toytree.tree(problem2_newick)
example2_tree.draw(use_edge_lengths=True, scalebar=True);

```

<!-- grade_id: problem2_newick, points: 1, grade: true, solution: false -->
```python
# tests for problem2_newick
problem2_tree = toytree.tree(problem2_newick) # this will raise an exception if the string is not valid
###
### AUTOGRADER TEST - DO NOT REMOVE
###
print("SUCCESS: problem2_newick passed all visible tests")
```

Output:
```
SUCCESS: problem2_newick passed all visible tests
```

## PROBLEM 3: Newick strings for all rooted trees with four leaves (1 POINT)
List all possible binary rooted trees with four leaves, $\{A, B, C, D\}$.  Assign a list containing the Newick strings for the trees to the variable `problem3_newick_list` below.  A rooted binary tree is tree where each node is either a leaf or has exactly two children.

```python
###
problem3_newick_list = ["((A,B),(C,D));",
                        "((A,C),(B,D));",
                        "((A,D),(B,C));",
                        "(A,(B,(C,D)));",
                        "(A,(C,(B,D)));",
                        "(A,(D,(B,C)));",
                        "(B,(A,(C,D)));",
                        "(B,(C,(A,D)));",
                        "(B,(D,(A,C)));",
                        "(C,(A,(B,D)));",
                        "(C,(B,(A,D)));",
                        "(C,(D,(A,B)));",
                        "(D,(A,(B,C)));",
                        "(D,(B,(A,C)));",
                        "(D,(C,(A,B)));"]
    
                                        
###

```

<!-- grade_id: problem3_newick_list, points: 1, grade: true, solution: false -->
```python
# tests for problem3_newick_list
assert isinstance(problem3_newick_list, list), "problem3_newick_list is not a list"

assert all(isinstance(list_element, str) for list_element in problem3_newick_list), "elements of problem3_newick_list are not all strings"

# the line below will raise an exception if the strings are not valid
problem3_trees = toytree.Multitree.MultiTree("\n".join(problem3_newick_list)) 

# check that each tree has the correct leaf set
for t in problem3_trees.treelist:
    assert set(t.get_tip_labels()) == {"A", "B", "C", "D"}, f'{t.write(tree_format=9)} does not have the correct leaf names'

# check that each tree is binary
for t in problem3_trees.treelist:
    assert t.is_bifurcating(), f'{t.write(tree_format=9)} is not a rooted binary tree'
    
# check for duplicate trees
def standardize_newick(s):
    t = toytree.tree(s)
    t.treenode.sort_descendants()
    return t.write(tree_format=9)

standardized_tree_set = set(map(standardize_newick, problem3_newick_list))
assert len(standardized_tree_set) == len(problem3_newick_list), "problem3_newick_list has duplicate tree structures"

###
### AUTOGRADER TEST - DO NOT REMOVE
###
print("SUCCESS: problem3_newick_list passed all visible tests")
```

Output:
```
SUCCESS: problem3_newick_list passed all visible tests
```

## PROBLEM 4:  A primate phylogenetic tree (1 POINT)
[Perelman et al.](https://doi.org/10.1371/journal.pgen.1001342) sequenced DNA from 186 primate species and constructed one of the most refined phylogenetic trees of the primates to date.  This phyogenetic tree is provided as the file `primates.newick` in this activity.  

Examine this tree (you will likely need to use large "height" and "width" arguments to the `tree.draw` method to see the tree clearly).  Are humans more closely related to [galagos](https://en.wikipedia.org/wiki/Galago), [tarsiers](https://en.wikipedia.org/wiki/Tarsier), or [lemurs](https://en.wikipedia.org/wiki/Lemur), according to this tree?  Give your answer by assigning the string "galagos", "tarsiers", or "lemurs" to the variable `humans_most_closely_related_to`.

```python
###
primates_tree_filename = "primates.newick"
primates_tree_newick = open(primates_tree_filename).read()
primates_tree = toytree.tree(primates_tree_newick)
primates_tree.draw(height=3000, width=2000)
humans_most_closely_related_to = "tarsiers"
###

```

<!-- grade_id: humans_most_closely_related_to, points: 1, grade: true, solution: false -->
```python
# tests for humans_most_closely_related_to
assert humans_most_closely_related_to in {"galagos", "tarsiers", "lemurs"}
###
### AUTOGRADER TEST - DO NOT REMOVE
###
print("SUCCESS: humans_most_closely_related_to passed all visible tests")
```

Output:
```
SUCCESS: humans_most_closely_related_to passed all visible tests
```

```python

```