# Day 15 notebook

The objectives of this notebook are to practice

* the UPGMA algorithm
* the neighbor joining algorithm

```python
import toytree
```

## The case of possible HIV transmission from dentist to patient
In the early 1990s a controversial case arose in which it was suspected that an HIV-positive dentist had transmitted HIV to at least one of his patients.  For some background on the case, you can read the [obituary of one of the patients who died of AIDS](https://www.nytimes.com/1991/12/09/obituaries/kimberly-bergalis-is-dead-at-23-symbol-of-debate-over-aids-tests.html).  Ultimately, the Centers for Disease Control and Prevention (CDC) became involved and obtained HIV samples from the dentist, HIV-positive patients of this dentist, and a number of HIV-positive individuals in the local community (which served as controls).  The CDC performed DNA sequencing on these samples and a subsequent phylogenetic analysis to determine whether or not the molecular data provided evidence that the dentist had transmitted HIV to his patients ([Ou et al. Science, 1992](http://science.sciencemag.org/content/256/5060/1165)).

We will revisit the data from this study and run the UPGMA and neighbor joining algorithms to see if we can reproduce its conclusions.  Below is a pairwise distance matrix for HIV genomic segments obtained from viral samples from five individuals: the dentist (D), patient A (PA), patient B (PB), local control 1 (C1), and local control 2 (C2).  These distances come from a multiple alignment of the V3 variable region of the HIV genome for these samples.

              C1      C2       D      PA      PB
      C1     0.0    0.09   0.098   0.105    0.12
      C2    0.09     0.0   0.072   0.076   0.101
       D   0.098   0.072     0.0    0.04   0.061
      PA   0.105   0.076    0.04     0.0   0.068
      PB    0.12   0.101   0.061   0.068     0.0

Here is that same matrix in the form of a Python dictionary-based distance matrix:

```python
v3_matrix = {
    ('C1', 'C1'): 0.0,   ('C1', 'C2'): 0.09,  ('C1', 'D'): 0.098, ('C1', 'PA'): 0.105, ('C1', 'PB'): 0.12,
    ('C2', 'C1'): 0.09,  ('C2', 'C2'): 0.0,   ('C2', 'D'): 0.072, ('C2', 'PA'): 0.076, ('C2', 'PB'): 0.101, 
     ('D', 'C1'): 0.098,  ('D', 'C2'): 0.072,  ('D', 'D'): 0.0,    ('D', 'PA'): 0.04,   ('D', 'PB'): 0.061,
    ('PA', 'C1'): 0.105, ('PA', 'C2'): 0.076, ('PA', 'D'): 0.04,  ('PA', 'PA'): 0.0,   ('PA', 'PB'): 0.068,
    ('PB', 'C1'): 0.12,  ('PB', 'C2'): 0.101, ('PB', 'D'): 0.061, ('PB', 'PA'): 0.068, ('PB', 'PB'): 0.0}
```

## PROBLEM 1: UPGMA (20 POINTS)
Run the UPGMA algorithm (by hand, or by code if you really want to) on this distance matrix.  Record your result by assigning the resulting tree, in Newick string format, to the variable `upgma_tree_newick` below.  Your tree should have branch lengths, with lengths rounded to three digits after the decimal point.  **Important note:** do *not* round the intermediate distances computed during the algorithm.

```python
###
upgma_tree_newick = '(C1:0.052,(C2:0.042,((D:0.02,PA:0.02):0.012,PB:0.032):0.009):0.01);'
###

```

```python
# draw your UPGMA tree
upgma_tree = toytree.tree(upgma_tree_newick)
upgma_tree.draw(use_edge_lengths=True, scalebar=True);
```

<!-- grade_id: upgma_tree_newick_valid_tree, points: 4, grade: true, solution: false -->
```python
# test upgma_tree_newick valid tree
upgma_tree = toytree.tree(upgma_tree_newick)
assert sorted(upgma_tree.get_tip_labels()) == ['C1', 'C2', 'D', 'PA', 'PB']
assert upgma_tree.is_rooted()
print("SUCCESS: upgma_tree_newick valid tree test passed")
```

Output:
```
SUCCESS: upgma_tree_newick valid tree test passed
```

<!-- grade_id: upgma_tree_newick_topology, points: 8, grade: true, solution: false -->
```python
# test upgma_tree_newick tree topology
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
# the code below extracts the branch lengths from your tree in a standardized order for testing
upgma_tree = toytree.tree(upgma_tree_newick)
upgma_tree.treenode.sort_descendants()
upgma_tree = toytree.tree(upgma_tree.write(tree_format=5)) # hack to make sure edges are in same order
upgma_branch_lengths = [round(x, 3) for x in upgma_tree.get_edge_values(feature="dist")]
print("The standardized list of branch lengths in your tree is:", upgma_branch_lengths, sep="\n")
print("The ith entry (0-based indexing) of this list will be tested with test upgma_tree_newick_branch_length_i")
```

Output:
```
The standardized list of branch lengths in your tree is:
[0.02, 0.02, 0.032, 0.042, 0.052, 0.012, 0.009, 0.01]
The ith entry (0-based indexing) of this list will be tested with test upgma_tree_newick_branch_length_i
```

<!-- grade_id: upgma_tree_newick_branch_length_0, points: 1, grade: true, solution: false -->
```python
# test upgma_tree_newick_branch_length_0
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: upgma_tree_newick_branch_length_1, points: 1, grade: true, solution: false -->
```python
# test upgma_tree_newick_branch_length_1
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: upgma_tree_newick_branch_length_2, points: 1, grade: true, solution: false -->
```python
# test upgma_tree_newick_branch_length_2
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: upgma_tree_newick_branch_length_3, points: 1, grade: true, solution: false -->
```python
# test upgma_tree_newick_branch_length_3
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: upgma_tree_newick_branch_length_4, points: 1, grade: true, solution: false -->
```python
# test upgma_tree_newick_branch_length_4
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: upgma_tree_newick_branch_length_5, points: 1, grade: true, solution: false -->
```python
# test upgma_tree_newick_branch_length_5
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: upgma_tree_newick_branch_length_6, points: 1, grade: true, solution: false -->
```python
# test upgma_tree_newick_branch_length_6
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: upgma_tree_newick_branch_length_7, points: 1, grade: true, solution: false -->
```python
# test upgma_tree_newick_branch_length_7
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

## PROBLEM 2: Neighbor joining (20 POINTS)
Run the neighbor joining algorithm (by hand, or by code if you really want to) on this distance matrix.  Record your result by assigning the resulting tree, in Newick string format, to the variable `nj_tree_newick` below.  Your tree should have branch lengths, with lengths rounded to three digits after the decimal point.  **Important note:** do *not* round the intermediate distances computed during the algorithm.

*Note: in Newick format, an unrooted tree has three children at the top level (e.g., "(A, B, C);")*

```python
###
nj_tree_newick = '(((C1:0.057,C2:0.033):0.022,PB:0.044):0.001,D:0.017,PA:0.023);'
###

```

```python
# draw your NJ tree
nj_tree = toytree.tree(nj_tree_newick)
nj_tree.draw(use_edge_lengths=True, scalebar=True);
```

<!-- grade_id: nj_tree_newick_valid_tree, points: 6, grade: true, solution: false -->
```python
# test nj_tree_newick valid tree
nj_tree = toytree.tree(nj_tree_newick)
assert sorted(nj_tree.get_tip_labels()) == ['C1', 'C2', 'D', 'PA', 'PB']
assert not nj_tree.is_rooted()
print("SUCCESS: nj_tree_newick valid tree test passed")
```

Output:
```
SUCCESS: nj_tree_newick valid tree test passed
```

<!-- grade_id: nj_tree_newick_topology, points: 7, grade: true, solution: false -->
```python
# test nj_tree_newick tree topology
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
# the code below extracts the branch lengths from your tree in a standardized order for testing
nj_tree = toytree.tree(nj_tree_newick)
rooted_nj_tree = nj_tree.root(names="C1")
rooted_nj_tree.treenode.sort_descendants()
standardized_nj_tree = toytree.tree(rooted_nj_tree.unroot().write(tree_format=5)) # hack to make sure edges are in same order
nj_branch_lengths = [round(x, 3) for x in standardized_nj_tree.get_edge_values(feature="dist")]
print("The standardized list of branch lengths in your tree is:", nj_branch_lengths, sep="\n")
print("The ith entry (0-based indexing) of this list will be tested with test nj_tree_newick_branch_length_i")
```

Output:
```
The standardized list of branch lengths in your tree is:
[0.023, 0.017, 0.044, 0.033, 0.057, 0.001, 0.022]
The ith entry (0-based indexing) of this list will be tested with test nj_tree_newick_branch_length_i
```

<!-- grade_id: nj_tree_newick_branch_length_0, points: 1, grade: true, solution: false -->
```python
# test nj_tree_newick_branch_length_0
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: nj_tree_newick_branch_length_1, points: 1, grade: true, solution: false -->
```python
# test nj_tree_newick_branch_length_1
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: nj_tree_newick_branch_length_2, points: 1, grade: true, solution: false -->
```python
# test nj_tree_newick_branch_length_2
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: nj_tree_newick_branch_length_3, points: 1, grade: true, solution: false -->
```python
# test nj_tree_newick_branch_length_3
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: nj_tree_newick_branch_length_4, points: 1, grade: true, solution: false -->
```python
# test nj_tree_newick_branch_length_4
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: nj_tree_newick_branch_length_5, points: 1, grade: true, solution: false -->
```python
# test nj_tree_newick_branch_length_5
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: nj_tree_newick_branch_length_6, points: 1, grade: true, solution: false -->
```python
# test nj_tree_newick_branch_length_6
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

## Follow-up analysis

Consider the following questions:
1. How do the two trees that you constructed differ, if at all?
2. Are these trees consistent with the possibility that the dentist transmitted HIV to the two patients considered here?

###
### Your answers here
###
